Effects of blood meal source and seasonality on reproductive traits of Culex quinquefasciatus (Diptera: Culicidae)

Author Response

The following is the authors’ response to the original reviews.

Reviewer #1 (Public Review):

Summary:

This study examines the role of host blood meal source, temperature, and photoperiod on the reproductive traits of Cx. quinquefasciatus, an important vector of numerous pathogens of medical importance. The host use pattern of Cx. quinquefasciatus is interesting in that it feeds on birds during spring and shifts to feeding on mammals towards fall. Various hypotheses have been proposed to explain the seasonal shift in host use in this species but have provided limited evidence. This study examines whether the shifting of host classes from birds to mammals towards autumn offers any reproductive advantages to Cx. quinquefasciatus in terms of enhanced fecundity, fertility, and hatchability of the offspring. The authors found no evidence of this, suggesting that alternate mechanisms may drive the seasonal shift in host use in Cx. quinquefasciatus.

Strengths:

Host blood meal source, temperature, and photoperiod were all examined together.

Weaknesses: The study was conducted in laboratory conditions with a local population of Cx. quinquefasciatus from Argentina. I'm not sure if there is any evidence for a seasonal shift in the host use pattern in Cx. quinquefasciatus populations from the southern latitudes.

We agree on the reviewers observation about the evidence on seasonal shift in the host use pattern in Cx. quinquefasciatus populations from southern latitudes. We include a paragraph in the Introduction section regarding this. Unfortunately, studies conducted in South America to understand host use by Culex mosquitoes are very limited, and there are virtually no studies on the seasonal feeding pattern. In Argentina, there is some evidence (Stein et al., 2013, Beranek, 2019) regarding the seasonal change in host use by Culex species, including Cx. quinquefasciatus, where the inclusion of mammals during the autumn has been observed. As part of a comprehensive study on characterising bridge vectors for SLE and WN viruses, our research group is currently working on the molecular identification of blood meals from engorged females to gain deeper insights into the seasonal feeding pattern of Culex mosquitoes. While the seasonal change in host use by Culex quinquefasciatus has not been reported in Argentina so far, there has been an observed increase in reported cases of SLE virus in humans between summer and fall (Spinsanti et al., 2008). It is based on this evidence that we hypothesise there is a seasonal change in host use by Cx. quinquefasciatus, similar to what occurs in the United States. This is also considering that both countries (Argentina and the United States) have regions with similar climatic conditions (temperate climates with thermal and hydrological seasonality). Since we work on the same species and in a similar temperate climate regimen, we assumed there is a seasonal shift in the host use by this mosquito species.

Reviewer #1 (Recommendations for the authors):

Abstract

Line 23: fed on two different hosts.

Accepted as suggested.

I think the concluding statement should be rewritten to say that immediate reproductive outcomes do not explain the shift in host use pattern of Cx. quinquefasciatus mosquitoes from birds to mammals towards autumn.

Accepted as suggested.

Introduction

No comments.

Materials and Methods

Please mention sample sizes in the text as well (n = ?) for each treatment.

Accepted as suggested.

Page 99: ......C. quinquefasciatus, since C. pipiens and its hybrids are present as well in Cordoba.

Accepted as suggested.

Results – Line 146: subsequently instead of posteriorly

Accepted all changes as suggested.

Line 148: were counted instead of was counted.

Accepted all changes as suggested.

Line 160: Subsequently instead of posteriorly

Accepted all changes as suggested.

Line 171: on fertility

Accepted all changes as suggested.

Line 174: there was an interaction effect on…

Accepted all changes as suggested.

Line 175: there were no differences in the number of eggs

Accepted all changes as suggested.

Discussion

I think the first paragraph in the discussion section is redundant and should be deleted.

The whole discussion was rewritten to be focused on our aims and results.

Line 282: this sentence needs to be rewritten.

Accepted as suggested.

Line 299: at 28{degree sign}C

Line 300: at 30{degree sign}C

Sorry, but we are not sure about your comment here. We checked. Temperatures are written as stated, 28°C and 30°C.

Line 363: I think the authors need to discuss more about the bigger question they were addressing. I think that the discussion section can be strengthened greatly by elaborating on whether there is evidence for a seasonal shift in host use pattern in Cx. quinquefasciatus in the southern latitudes. If yes, what alternate mechanisms they believe could be driving the seasonal change in host use in this species in the southern latitudes now that they show the 'deriving reproductive advantages' hypothesis to be not true for those populations.

Thanks for this observation. We agree and so the Discussion section was restructured to align it with our results, as suggested.

Reviewer #2 (Public Review):

Summary:

Conceptually, this study is interesting and is the first attempt to account for the potentially interactive effects of seasonality and blood source on mosquito fitness, which the authors frame as a possible explanation for previously observed host-switching of Culex quinquefasciatus from birds to mammals in the fall. The authors hypothesize that if changes in fitness by blood source change between seasons, higher fitness in birds in the summer and on mammals in the autumn could drive observed host switching. To test this, the authors fed individuals from a colony of Cx. quinquefasciatus on chickens (bird model) and mice (mammal model) and subjected each of these two groups to two different environmental conditions reflecting the high and low temperatures and photoperiod experienced in summer and autumn in Córdoba, Argentina (aka seasonality). They measured fecundity, fertility, and hatchability over two gonotrophic cycles. The authors then used a generalized linear mixed model to evaluate the impact of host species, seasonality, and gonotrophic cycle on fecundity and fertility and a null model analysis via data randomization for hatchability. The authors were trying to test their hypothesis by determining whether there was an interactive effect of season and host species on mosquito fitness. This is an interesting hypothesis; if it had been supported, it would provide support for a new mechanism driving host switching. While the authors did report an interactive impact of seasonality and host species, the directionality of the effect was the opposite of that hypothesized. While this finding is interesting and worth reporting, there are significant issues with the experimental design and the conclusions that are drawn from the results, which are described below. These issues should be addressed to make the findings trustworthy.

Strengths:

(1) Using a combination of laboratory feedings and incubators to simulate seasonal environmental conditions is a good, controlled way to assess the potentially interactive impact of host species and seasonality on the fitness of Culex quinquefasciatus in the lab.

(2) The driving hypothesis is an interesting and creative way to think about a potential driver of host switching observed in the field.

Weaknesses:

(1) There is no replication built into this study. Egg lay is a highly variable trait, even within treatments, so it is important to see replication of the effects of treatment across multiple discrete replicates. It is standard practice to replicate mosquito fitness experiments for this reason. Furthermore, the sample size was particularly small for some groups (e.g. 15 egg rafts for the second gonotrophic cycle of mice in the autumn, which was the only group for which a decrease in fecundity and fertility was detected between 1st and 2nd gonotrophic cycles). Replicates also allow investigators to change around other variables that might impact the results for unknown reasons; for example, the incubators used for fall/summer conditions can be swapped, ensuring that the observed effects are not artefacts of other differences between treatments. While most groups had robust sample sizes, I do not trust the replicability of the results without experimental replication within the study.

We agree egg lay is a variable trait and so we consider high numbers of mosquitoes and egg lay during experiments compared to our studies of the same topics. Evaluating variables such as fecundity, fertility, or other types of variables (collectively referred to as "life tables") is a challenging issue that depends on several intrinsic and extrinsic factors. Because all of this, in some experiments, sample sizes might not be very large, and in several articles, lower sample sizes could be found. For instance, in Richards et al. (2012), for Culex quinquefasciatus, during the second gonotrophic cycle, some experiments had 13 or even 6 egg rafts. For species like Aedes aegypti, the sample size for life table analysis is also usually small. As an example, Muttis et al. (2018) reported between 1 and 4 engorged females (without replicates). In addition, small sample size would be a problem if we would not have obtained any effect, which is not the case due to the fact that we were interested in finding an effect, regardless of the effect size. Because of this, we do find our sample sizes quite robust for our results.

Regarding the need to repeat the experiments in order to give more robustness to the study we also agree. However, after a review of the literature (articles cited in the original manuscript), it is apparent that similar experiments are not frequently repeated as such. Examples of this are the studies of Richards et al. (2012), Demirci et al. (2014) or Telang & Skinner (2019), which even they manipulate several cages at a time as “replicates”, they are not true replicates because they summarise and manipulate all data together, and do not repeat the experiment several times. We see these “replicates” as a way of getting a greater N.

As was stated by the reviewer, repetition is a resource and time-consuming activity that we are not able to do. Replicating the experiment poses a significant time and resources challenge. The original experiment took over three months to complete, and it is anticipated that a similar timeframe would be necessary for each replication (6 months in total considering two more replicates). Given our existing commitments and obligations, dedicating such an extensive period solely to this would impede progress on other crucial projects and responsibilities.

Given the limitations of resources and time and the infrequent use of experimental replication in this type of studies, we performed a simulation-based analysis via a Monte Carlo approach. This approach involved generating synthetic data that mimics the expected characteristics of the original experiment and subsequently subjecting it to the same analysis routine. The main goal of this simulation was to evaluate the potential spuriousness and randomness of the results that might arise due to the experimental conditions. So, evaluating the robustness and confidence of our results and data.

(2) Considering the hypothesis is driven by the host switching observed in the field, this phenomenon is discussed very little. I do not believe Cx. quinquefasciatus host switching has been observed in Argentina, only in the northern hemisphere, so it is possible that the species could have an entirely different ecology in Argentina. It would have been helpful to conduct a blood meal analysis prior to this experiment to determine whether using an Argentinian population was appropriate to assess this question. If the Argentinian populations don't experience host switching, then an Argentinian colony would not be the appropriate colony to use to assess this question. Given that this experiment has already been conducted with this population, this possibility should at least be acknowledged in the discussion. Or if a study showing host switching in Argentina has been conducted, it would be helpful to highlight this in the introduction and discussion.

Thanks for this observation. We agree. However, we conducted the experiment beside host use data from Argentina since we used the mosquito species, and the centre region of Argentina (Córdoba) has a similar temperate weather regimen that those observed in the east coast of US.

We are aware that few studies regarding host shifting in South America are available, some such that those conducted by Stein et al. (2013) and Beranek (2019) reported a moderate host switch for Culex quinquefasciatus in Argentina. We have already performed a study about seasonal host feeding patterns for this species. However, even though there are few studies regarding host shifting, our hypothesis is based mainly in the seasonality of human cases of WNV and SLEV, a pattern that has been demonstrated for our region, see for example the study of Spinsanti et al. (2008).

We include a new paragraph in the Introduction and Discussion sections. Please see answers Reviewer #1.

(3) The impacts of certain experimental design decisions are not acknowledged in the manuscript and warrant discussion. For example, the larvae were reared under the same conditions to ensure adults of similar sizes and development timing, but this also prevents mechanisms of action that could occur as a result of seasonality experienced by mothers, eggs, and larvae.

We understand the confusion that may have arisen due to a lack of further details in the methodology. If we are not mistaken, you are referring to our oversight regarding the consideration of carry-over effects of larvae rearing that could potentially impact reproductive traits. When investigating the effects of temperature or other environmental factors on reproductive traits, it is possible to acclimate either larvae or adults. This is due to the significant phenotypic plasticity that mosquitoes exhibit throughout their entire ontogenetic cycle. In our study, we followed an approach similar to that of other authors where the adults are exposed to experimental conditions (temperature and photoperiod). For a similar approach you can refer to the studies conducted by Ferguson et al. (2018) for Cx. pipiens, Garcia Garcia & Londoño Benavides (2007) for Cx. quinquefasciatus or Christiansen-Jucht et al. (2014, 2015) for Anopheles gambiae.

(4) There are aspects of the data analysis that are not fully explained and should be further clarified. For example, there is no explanation of how the levels of categorical variables were compared.

The methodology and statistical analysis were expanded for a better understanding.

(5) The results show the opposite trend as was predicted by the authors based on observed feeding switches from birds to mammals in the autumn. However, they only state this once at the end of the discussion and never address why they might have observed the opposite trend as was hypothesized.

The discussion was restructured to focus on our results and our model.

(6) Generally speaking, the discussion has information that isn't directly related to the results and/or is too detailed in certain parts. Meanwhile, it doesn't dig into the meaning of the results or the ways in which the experimental design could have influenced results.

As mentioned above, the discussion was restructured to reflect our findings. We also included the effect that our design might have influenced our results. However, as stated above we do not fully agree that the design is inadequate for our analysis, we performed standard protocols followed by other researchers and studies in this research field.

(7) Beyond the issue of lack of replication limiting trust in the conclusions in general, there is one conclusion reached at the end of the discussion that would not be supported, even if additional replicates are conducted. The results do not show that physiological changes in mosquitoes trigger the selection of new hosts. Host selection is never measured, so this claim cannot be made. The results don't even suggest that fitness might trigger selection because the results show that physiological changes are in the opposite direction as what would be hypothesized to produce observed host switches. Similarly, the last sentence of the abstract is not supported by the results.

We agree with this observation. However, we did not evaluate the impact of fitness on host selection in this study. Instead, we aimed to investigate the potential influence of seasonality on mosquito fitness as a potential trigger for a shift in host selection. We agree that we have incorrectly used the term “host selection” when we should actually be discussing “host use change”. Our results indicate a seasonal alteration in mosquito fitness in response to temperature and photoperiod changes. Building upon this observation, we re-discussed our hypothesis and theoretical model to explain this seasonal shift in host use.

(8) Throughout the manuscript, there are grammatical errors that make it difficult to understand certain sentences, especially for the results.

All English grammar and writing of the manuscript was revised and corrected to be easily understood.

This study is driven by an interesting question and has the potential to be a valuable contribution to the literature.

Reviewer #2 (Recommendations for The Authors):

I hope that the authors will consider the suggested revisions and experimental replication to improve the quality of the study and paper.

This study tests a very interesting hypothesis. I understand that additional replicates are difficult to conduct, but I do believe that fitness studies absolutely require experimental replicates. Unless you are able to replicate the observed effects, I personally would not trust the results of this study. I hope that you will consider conducting replicates so that this important question can be answered in a more robust manner. Below, I expand upon some additional points in the public review and also provide more specific suggestions. I provided some copy-editing feedback, but was not able to point out all grammatical mistakes. I suggest that you use ChatGPT to help you edit the English. For example, you can feed ChatGPT your MS and ask it to bold the grammatical errors or you can ask it to edit grammatical errors and bold the sections that were edited. I understand that writing in a second language is very difficult (from personal experience!), so I view ChatGPT as a great tool to help even the playing field for publishing. Below are line item suggestions. Apologies that wording is curt, I was trying to be efficient in writing.

20-21: I suggest that you emphasize that you are investigating the interactive effect.

Accepted as suggested.

22: they weren't "reared" (from larvae) in different conditions, they were "maintained" as adults

Accepted as suggested.

26-27: increased/decreased is a bit misleading since you did not evaluate these groups sequentially in time. It might be more accurate to describe it as less than/greater than. Also, if you say increased/decreased or less than/greater than, you should always say what you are comparing to. The same applies throughout the MS.

Accepted as suggested.

29-30: "finding the" is not correct here; could be "with the lowest..."

Accepted as suggested.

34-36: I do not think that your results suggest this, even if you were to replicate the results of this experiment. You haven't shown metabolic changes.

We understand the point. Accepted as suggested.

42-44: "one of the main responsible" should be "one of the main species responsible..."

Accepted as suggested.

48: I think that "host preference" is better than selection here; -philic denotes preference

Accepted as suggested.

50: "Moreover" isn't the correct transition word here

Accepted as suggested.

57: "could" isn't correct here; consider saying "... species sometimes feed primarily on mammal hosts, including humans, in certain situations."

Accepted as suggested.

58: Different isn't correct word here

Accepted as suggested.

60: delete "feeding"

Accepted as suggested.

66-68: I am not familiar with any blood meal analysis studies in the southern hemisphere that show host switching for Culex species between summer and autumn. If this hasn't been shown, then this critique of the host migration hypothesis doesn't make sense.

There are some studies pointing this out (Stein et al., 2013, Beranek 2019), and unpublished data from us). However, our hypothesis has supported by epidemiological data observed in human population which indicate a seasonal activity pattern. It was explained in depth in the Introduction section.

68: ensures is not the right word; I suggest "suggests"

Accepted as suggested.

68-70: this explanation isn't clear to me; please revise

It will be revised. Accepted as suggested.

70: change cares to care

Accepted as suggested.

76-77: can you explain how they were not supported by the data for the benefit of those who are not familiar with these papers please?

Accepted as suggested.

87-89: I suggest the following wording: "In the autumn, we expect a greater number of eggs (fecundity) and larvae (fertility) in mosquitoes after feeding on a mammal host compared to an avian host, and the opposite relationship in the summer."

Accepted as suggested.

99: edit for grammar

Accepted as suggested.

102: suggest: "...offered a blood meal from a restrained chicken twice a month"

Accepted as suggested.

107: powder

Accepted as suggested.

108: inbred? Is this the term you meant to use?

Changed as suggested.

109: "several" cannot be used to describe 20 generations; suggest using "over twenty generations"; also, it would be good to acknowledge in your discussion that lab adaptation could force evolution, especially since mosquitoes are kept at constant temperatures and fed with certain hosts (with easy access) in the lab. Also, it would be good to know when the experiments were conducted to know the lapse of time between the creation of the colony and the experiments.

Accepted as suggested.

110-111: Does humidity vary between summer and fall in Córdoba? If so, I suggest acknowledging in the discussion that if humidity differences are involved in a potential interaction between host species and seasonality, then this would not have been captured by your experimental design.

Several variables change during seasons. We were interested in capturing the effects of temperature and photoperiod, since humidity is a variable difficult to control.

113-116: I suggest combining into one sentence to make more concise.

Accepted as suggested.

135: You might be obscuring the true impact of seasonality by rearing the larvae under the same conditions. There may be signals that mothers/eggs/larvae receive that influence their behavior (e.g. I believe this is the case for diapause), so this limitation should also be acknowledged. I understand why you decided to do this to control for development time and size, but it is something that should be considered in the discussion.

As it was explained above, Cx. quinquefasciatus do not suffer diapause in our country. Maintaining mosquitoes from adults was an approach selected by us based on other studies.

138: edit: "with cotton pads soaked in... on plastic..."; what is plastic glass? Do you mean plastic dishes?

Accepted as suggested.

141: here and throughout paragraph, full should be "fully"

Accepted as suggested.

144: located should be "placed"

Accepted as suggested.

147: suggest editing to "at which point, they were fixed with 1 mL of 96% ethanol and the number of L1 larvae per raft was counted."

Accepted as suggested.

154-155: edit for grammar

Accepted as suggested.

157: Your GLM explanation doesn't say anything about how you made pairwise comparisons between your levels; did you use emmeans?

This revised version includes a more detailed methodology and statistical analysis. Accepted as suggested.

158-160: I don't understand why you took this approach - it seems strange to me to use this analysis, but I am not familiar with it, so it might be that I lack the knowledge to be able to adequately evaluate. Please provide more explanation so that readers can better understand this analysis. A citation for this kind of application of the analysis would be helpful.

It was changed to be in accordance with the remaining analyses.

173: replace neither with either

Accepted as suggested.

174: this applies throughout; edit to : "An interaction effect was observed..."

Accepted as suggested.

175: "it was not found" is grammatically incorrect; instead : "We did not find ..." or "no differences in... were detected", etc

Accepted as suggested.

183: "it was detected" is grammatically incorrect

Accepted as suggested.

185-186: "being this treatment... in terms of fitness": I do not understand what this means. Please rephrase

Accepted as suggested.

170-199: you should provide the effect sizes and p values in text and/or in the figure for the pairwise comparisons

Accepted as suggested.

193-196. These two sentences are confusing and I am not sure what you mean, especially in the first sentence.

It was rewritten. Accepted as suggested.

Figure 1: This figure is great and easy to read and interpret! Thank you for the comment! 218-219: it is important to state which mosquito species you are referring to here.

Accepted as suggested.

226-227: you definitely should acknowledge the small sample size here.

Considered.

227: "it was observed" should be "We observed" or "A greater hatching rate.... was observed."

Accepted as suggested.

228-229: is the result really comparable even though you took very different approaches to the analysis for these outcomes?

Changed to be comparable.

230-278: the discussion of these hypotheses is too long and detailed, especially since the comparison of mouse vs chicken wasn't your main question; you really wanted to understand this in the context of seasonality. I suggest cutting this down a lot and making room to dig into your results more, and also to discuss the potential impacts of your experimental design/limitations on the results.

Discussion was changed to focus on our results and model. Accepted as suggested.

281: Hoffman is an old citation; I suggest you cite a modern review.

Accepted as suggested. We deleted it due to the re-writing of the manuscript.

282: "It can be recognise".. I am not sure what you are trying to say here

Accepted as suggested.

1. After the first time you write a species name, you can abbreviate the genus in all future mentions unless it is at the beginning of a sentence.

Accepted as suggested.

303-305: Revise this sentence. E.g "Fewer studies are available regarding photoperiod and show mixed results; Mogi (1992) found that mid and long day lengths induced greater fecundity while Costanzo et al. (2015) did not find differences in fecundity by day length."

Accepted as suggested.

315-316: typically, unpublished data shouldn't be referenced; I'm not sure if eLife has a policy on this.

We will check this with eLife guidelines. However, since the lack of evidence on this pattern we consider important to include this unpublished data.

316: Aegypti should be lowercase

Accepted as suggested.

328-330: This sentence is redundant with the first sentence of the paragraph

Accepted as suggested.

321-336: You never reintroduced your hypothesis in your discussion. I suggest that you center your whole discussion more directly around the hypothesis that motivated the study. If you decide not to restructure your discussion, you should at least reintroduce your hypothesis here and discuss how your results do not support the hypothesis.

Accepted as suggested.

337-348: This paragraph is a bit confusing as you jump between fertility and hatchability

Accepted as suggested.

353: is viral transmission the right word to use here? I think you might mean bridge vector transmission to humans specifically?

Accepted as suggested.

357: you say "neither" but never define which traits you are referring to

Accepted as suggested.

361: I suggest "two variables previously analyzed separately..."

Accepted as suggested.

General: There is no statement about the availability of data; it is eLife policy to require all data to be publicly available. Also, it would be helpful to share your code to help understand how you conducted pairwise comparisons, etc.

In the submission it was not mentioned anything about data availability. However, all data and scripts will be uploaded with the VOR if it is required.

Recommendations for the authors:

I found your study interesting and potentially promising. However, there are some fundamental problems with the study design and the hypothesis, including:

<(1) Seasonality simulation - Seasonality is strongly associated with time, so it is unusual to simulate seasonal factors without accounting for time. The actual factors associated with seasonal change in reproductive output may be neither a difference in host blood meal nor temperature and photoperiod. It is therefore, odd to reduce seasonality to a difference in photoperiod and temperature in summer and autumn without even mentioning the time of year when the experiment was carried (except for the mention of February as the time the stock samples were collected from the wild).

The temperature and photoperiod settings are established according to a representative day in both autumn and summer. To determine these settings, we utilized climate data spanning a 3-year period (2020-2022), encompassing the most frequently occurring temperatures and day lengths. The weather conditions remained notably consistent throughout this time frame, which is why the specific year was not mentioned. Moreover, including the year in laboratory experiment details is uncommon, as evident in various papers. This practice can be corroborated by referring to multiple sources (cited in the original manuscript). We mention this in the new version.

(2) Hypothesis - While the hypothesis alludes to the 'reason' for seasonal host shift, the prediction is on the outcome of the interaction between blood meal type and season.

It might be nicer to frame your hypothesis to be consistent with the aim, which is, testing the partial contributions of blood meal type, versus photoperiod and temperature to seasonal change in the reproductive output of Culex quinquefasciatus. A hypothesis like that can be accompanied by alternative predictions according to the expected individual and interactive effects of both factors.

It was rewritten in the revised version to be consistent with our predictions and findings.

Blood meal type, temperature, and photoperiod are all components of seasonality, so the strength of the study is its potential to decouple the effect of blood meal type from that of temperature and photoperiod on the seasonal reproductive output of Culex quinquefasciatus by comparing the two blood meal types under simulated summer and winter conditions. Ideally, this should have been over a natural summer and winter because a natural time difference captures the effect of other seasonal factors other than temperature and photoperiod.

Furthermore, the hypothesis stemmed from field observations, while the study itself was conducted under laboratory conditions using a local population of Culex quinquefasciatus from Argentina. It remains uncertain whether there is supporting evidence for a seasonal shift in host usage in Culex quinquefasciatus from the stock population. Discussing the field observations within the stock population would provide valuable insights.

It was considered in the new version.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Host selection by mosquitoes is a keystone to understand viral circulation and predict future infection outbreaks. Culex mosquitoes frequently feed on birds during spring and early summer, shifting into mammals towards late summer and autumn. This host switch may be produced by changes in mosquito fitness. The aim of this study was to assess if blood meal source and seasonality may influence reproductive traits of Culex quinquefasciatus mosquitoes. For this purpose, Cx. quinquefasciatus mosquitoes were reared in simulated summer and autumn conditions and fed with two different hosts, chickens and mice, in a factorial design. Fecundity, fertility and hatchability during two consecutive gonotrophic cycles were estimated. We found a greater fecundity and fertility for mosquitoes fed upon birds than mammals. Fecundity and fertility increased in autumn for chicken-fed mosquitoes, whereas they decreased for mouse-fed mosquitoes. These traits decreased in the second gonotrophic cycle for mouse-fed mosquitoes, whereas they did not vary between cycles for chicken-fed mosquitoes. Blood meal source had a significant effect on hatchability, finding the lowest hatchability in mouse-fed mosquitoes during the second gonotrophic cycle, whereas no differences were detected among the remaining treatments. These results indicate that fecundity and fertility are influenced by blood meal source and seasonality, increasing towards autumn for mosquitoes fed on birds and decreasing for mosquitoes fed on mammals. Hatchability seems also being influenced by blood meal source. Our results suggest that blood meal source and seasonality induce metabolic changes in the mosquito reproductive physiology which might affect host selection patterns.

Revised version: Host selection by mosquitoes is a keystone to understanding viral circulation and predicting future infection outbreaks. Culex mosquitoes frequently feed on birds during spring and early summer, shifting into mammals towards late summer and autumn. This host switch may be due to changes in mosquito fitness. The aim of this study was to assess if the interaction effect of blood meal source and seasonality may influence reproductive traits of Culex quinquefasciatus mosquitoes. For this purpose, Cx. quinquefasciatus mosquitoes were reared in simulated summer and autumn conditions and fed on two different hosts, chickens and mice, in a factorial design. Fecundity, fertility, and hatchability during two consecutive gonotrophic cycles were estimated. We found greater fecundity and fertility for mosquitoes fed upon birds than mammals. Fecundity and fertility increased in autumn for chicken-fed mosquitoes, whereas they decreased for mouse-fed mosquitoes. These traits decreased in the second gonotrophic cycle for mouse-fed mosquitoes, whereas they did not vary between cycles for chicken-fed mosquitoes. The effect of hatchability on fertility was rather limited. These results indicate a statistically significant interaction effect of blood meal source and seasonality on fecundity and fertility. However, the pattern was opposite in relation to our hypothesis, suggesting that further studies are needed to confirm and expand our knowledge about mosquito biology and its relationship with seasonal host use shifting.

Introduction

Earlier version: The southern house mosquito, Culex quinquefasciatus Say (Diptera: Culicidae), is a worldwide distributed vector of several pathogens, such as Wuchereria bancrofti Cobbold (Kasili et al., 2009), Dirofilaria immitis Leidy (Labarthe et al., 1998) and avian malaria (Farajollahi et al., 2011). Besides, this mosquito is also one of the main responsible for the transmission of the arboviruses West Nile virus (WNV) and Saint Louis encephalitis virus (SLEV), among others (Farajollahi et al., 2011). Because of this, this species is of critical concern due to its impact on both public and veterinary health.

Revised version: The southern house mosquito, Culex quinquefasciatus Say (Diptera: Culicidae), is a worldwide distributed vector of several pathogens, such as Wuchereria bancrofti Cobbold (Kasili et al., 2009), Dirofilaria immitis Leidy (Labarthe et al., 1998) and avian malaria (Farajollahi et al., 2011). Besides, this mosquito is also one of the main species responsible for the transmission of the arboviruses West Nile virus (WNV) and Saint Louis encephalitis virus (SLEV), among others (Farajollahi et al., 2011). Because of this, this species is of critical concern due to its impact on both public and veterinary health.

Earlier version: Since arboviruses are transmitted by the bite of infectious mosquitoes between vertebrate hosts, activity of these pathogens is determined, at least in part, by the blood feeding habits of these vectors. Host selection of Culex mosquitoes is variable and it can be classified in four categories: (1) mammophilic (Muturi et al., 2008), (2) ornithophilic (Jansen et al., 2009), (3) herpetophilic (Janssen et al., 2015) and (4) generalist (Molaei et al., 2007). Moreover, some species may experience a host switch in its feeding habits, from avian to mammal hosts (Hancock & Camp, 2022). These species are of greater concern from an epidemiological standpoint because they can act as bridge vectors in the transmission of some arboviruses, such as SLEV and WNV (Kilpatrick et al., 2005).

Revised version: Since arboviruses are transmitted by the bite of infected mosquitoes between vertebrate hosts, activity of these pathogens is determined, at least in part, by the blood feeding habits of these vectors. Host preference of Culex mosquitoes is variable and it can be classified in four categories: (1) mammophilic (Muturi et al., 2008), (2) ornithophilic (Jansen et al., 2009), (3) herpetophilic (Janssen et al., 2015) and (4) generalist (Molaei et al., 2007). Additionally, some species may experience a host switch in its feeding habits, from avian to mammal hosts (Hancock & Camp, 2022). These species are of greater concern from an epidemiological standpoint because they can act as bridge vectors in the transmission of some arboviruses, such as SLEV and WNV (Kilpatrick et al., 2005).

Earlier version: Culex quinquefasciatus is considered an ornithophilic species in several areas around the globe (Takken & Verhulst, 2013). Despite its ornithophilic nature, Cx. quinquefasciatus and other Culex species could feed primarily from mammal hosts, including the human, in certain situations. Different studies have pointed out that in autumn there is an increased activity of SLEV and WNV in human populations (Kilpatrick et al., 2006, Spinsanti et al., 2008), suggesting that viral activity spills over to mammals from birds as a result of the mosquito feeding host switch (Edman & Taylor, 1968, Kilpatrick et al., 2006). A few hypotheses have attempted to explain why primarily ornithophilic mosquitoes also feed upon mammals in some contexts. One of the most widely accepted hypothesis states that the shift on feeding behavior of Culex mosquitoes is a consequence of the autumn migration of American robin (Turdus migratorius), the principal host in the northeastern US, generating an opportunistic change into mammal hosts (Kilpatrick et al., 2006). However, because this avian species is lacking in Southern latitudes it is an implausible hypothesis for many austral areas, being restricted only for some regions of the US. A second hypothesis proposed by Burkett-Cadena et al. (2011) ensures that host breeding cycles drive the host shift of mosquitoes, and peaks of host use are detected in periods of reproductive investment. According to this hypothesis, in hosts with parental cares, during reproductive seasons there is a greater investment of energy in assuring the offspring survivorship, consequently producing an increase in susceptibility of being bitten by mosquitoes as a result of a decrease in defensive behaviors (Burkett-Cadena et al., 2011). Other hypotheses suggest that mosquito host shift is a result of changes in host preference throughout the seasons (Edman, 1974), changes in density and feeding success of mosquitoes (Nelson et al., 1976) or seasonal changes in host habitat use by mosquitoes, however, they were poorly or not supported by field data (Edman, 1974; Burkett-Cadena et al., 2011).

Revised version: Culex quinquefasciatus is generally considered an ornithophilic species in several areas around the globe (Takken & Verhulst, 2013). Despite its ornithophilic nature, this and other Culex species sometimes feed primarily on mammal hosts, including humans, in certain situations. Various studies have pointed out that in autumn there is an increased activity of SLEV and WNV in human populations (Kilpatrick et al., 2006; Spinsanti et al., 2008), suggesting that viral activity spills over to mammals from birds as a result of the mosquito host switch (Edman & Taylor, 1968; Kilpatrick et al., 2006). A few hypotheses have attempted to explain why primarily ornithophilic mosquitoes also feed upon mammals in some contexts. One of the most widely accepted hypothesis states that the shift on feeding behaviour of Culex mosquitoes is a consequence of the autumn migration of American robin (Turdus migratorius), the principal host in the northeastern US, generating an opportunistic change into mammal hosts (Kilpatrick et al., 2006). In South America in general and in Argentina in particular, studies regarding host switch are scarce. Nevertheless, a few limited studies have indicated a seasonal variation throughout the seasons (Stein et al., 2013; Beranek, 2018). Despite that, Argentina has acknowledged seasonal variation in the activity of SLEV in humans, which can be attributed to variations in mosquito host-feeding pattern (Spinsanti et al., 2008). In Southern latitudes, species such as T. migratorius are lacking, hence the “migration hypothesis” is very unlikely for many austral areas, being restricted only for some regions of the US.

Added in the revised version: A second hypothesis proposed by Burkett-Cadena et al. (2011) suggests that host breeding cycles drive the host shift of mosquitoes. According to this hypothesis, in hosts with parental care, during reproductive seasons there is a greater investment of energy in assuring the offspring survivorship, consequently producing an increase in susceptibility of being bitten by mosquitoes as a result of a decrease in defensive behaviours. This event leads to the detection of peaks of host use during periods of reproductive investment, in summer for birds and autumn for mammals (Burkett-Cadena et al., 2011).

Earlier version: Several biological (stress, metabolic rate, blood meal source) and environmental (temperature, photoperiod, i.e., season) variables that influence mosquito physiology might have an effect on reproduction traits, such as fecundity, development rate and survivorship, generating new nutritional requirements (Ciota et al., 2014, Costanzo et al., 2015, Gervasi et al., 2016; Yan et al. 2017, 2018). These factors, ultimately, could produce seasonal variation in host use by mosquitoes. The aim of the present study is to evaluate the effect of both host blood meal source and seasonality (combination of temperature and photoperiod) on three reproductive traits (fecundity, fertility and hatchability) of Culex quinquefasciatus mosquitoes. We hypothesize that the seasonal host switch of mosquitoes is the result of the effect of blood meal source and seasonality on reproductive outputs. We expect a greater number of eggs (fecundity) and larvae (fertility) in mosquitoes fed on mammal in autumn than in an avian host, expecting the opposite pattern in summer.

Revised version: While some other minor hypotheses, unsupported by field data, have been proposed, none of them have been specifically focused on vector biology. Numerous biological factors, including stress, metabolic rate, and the source of a blood meal, as well as environmental variables such as temperature and photoperiod, exert an influence on mosquito physiology that might affect various reproductive traits such as fecundity, development rate or survivorship. This set of factors gives rise to new nutritional requirements that, ultimately, may lead to seasonal variation in host choice by mosquitoes (Ciota et al., 2014; Costanzo et al., 2015; Gervasi et al., 2016; Yan et al. 2017, 2018). In the literature, several studies have evaluated the impact of these biological and environmental variables on mosquito reproduction, as mentioned before. However, due to the complex nature of mosquito biology, there are likely multiple interactions among variables that give rise to new responses that may not be observed individually. Therefore, the aim of the present study was to assess whether exists an interaction effect between the source of the host blood meal and seasonality (in terms of temperature and photoperiod) on three reproductive traits of Culex quinquefasciatus mosquitoes, fecundity, fertility, and hatchability. Our hypothesis states that the interaction between these two variables influence the reproductive outcomes, potentially leading to a seasonal shift in host selection, driven by a reproductive advantage. Considering the reported host preference changes in Cx. quinquefasciatus, in autumn we expect a greater number of eggs (fecundity) and larvae (fertility) in mosquitoes after feeding on a mammal host compared to an avian host, and the opposite relationship in summer.

Materials and methods

Added in the revised version: Establishment and maintenance of mosquitoes

Removed in the revised version: Establishment and rearing of mosquito colony

Earlier version: Egg rafts of Culex quinquefasciatus were collected from a drainage ditch at Universidad Nacional de Córdoba Campus, Córdoba city in February 2021. Each raft was maintained individually in plastic containers with one liter of distilled water. The hatched larvae were fed with 100 mg of liver power three times per week until pupation. Pupae were transferred to a plastic emerging cage (21 cm x 12 cm) covered with a tulle-like fabric, containing distilled water and without food. Adults emerging from each raft were identified morphologically (Darsie, 1985) and molecularly (Smith & Fonseca, 2004) to make sure they corresponded to Cx. quinquefasciatus since in Córdoba is also present Cx. pipiens and its hybrids (Branda et al., 2021). All adults were reared in a cardboard cage of 22.5 liters (28 cm x 36.5 cm) and fed ad libitum with a 10% of sugar solution soaked in cotton pads on plastic glasses. For long-term maintenance of the colony, 24-h starved mosquitoes were offered with a blood meal of a restrained chicken twice a month. Four days post-feeding, a plastic recipient with distilled water was located inside the cage to allow engorged females lay egg rafts. Batches were collected and transferred to plastic containers (30 cm x 25 cm x 7 cm) in a proportion of 3 rafts per container, filled with 3 liters of distilled water. The hatched larvae were also fed with liver power at the same proportion described above. Pupae were transferred to emerging cages and adults to the final cardboard cage. The colony was inbred and maintained during several generations (generation > F20) in the Insectary of the Instituto de Virología “Dr. J.M. Vanella” (InViV). Room controlled conditions were 28°C, 12L:12D photoperiod and 70% relative humidity.

Revised version: Egg rafts of Culex quinquefasciatus were collected from a drainage ditch at Universidad Nacional de Córdoba Campus, Córdoba city, in February 2021. Each raft was individually maintained in plastic containers with one liter of distilled water. The hatched larvae were fed with 100 mg of liver powder three times per week until pupation. Pupae were then transferred to plastic emerging cages (21 cm x 12 cm) covered with a tulle-like fabric, containing distilled water but no food. Adults emerging from each raft were identified morphologically (Darsie, 1985) and molecularly (Smith & Fonseca, 2004) to ensure they corresponded to Cx. quinquefasciatus, since Cx. pipiens and its hybrids are also present in Córdoba (Branda et al., 2021). All adults were reared in a cardboard cage of 22.5 liters (28 cm x 36.5 cm) and provided ad libitum with a 10% sugar solution soaked in cotton pads placed on plastic cups. For long-term maintenance of the colony, 24-hour-starved mosquitoes were offered a blood meal from a restrained chicken twice a month. Four days after feeding, a plastic container with distilled water was placed inside the cage to allow engorged females to lay egg rafts. Batches of egg were collected and transferred to plastic containers (30 cm x 25 cm x 7 cm) in a proportion of 3 rafts per container, filled with three liters of distilled water. The hatched larvae were also fed with liver powder at the same proportion described above. Pupae were transferred to emerging cages, and adult mosquitoes were placed in the final cardboard cage. The colony has been maintained for over twenty generations in the Insectary of the Instituto de Virología “Dr. J.M. Vanella” (InViV). Room conditions were controlled at 28°C, with a 12L:12D photoperiod and 70% relative humidity.

Experimental design

Blood source

Earlier version: For experimental trials, an avian and mammal model was used to evaluate the effect of blood meal source on fecundity, fertility and hatchability. The former host was represented by live chicks from the species Gallus gallus and the latter by live mice of the species Mus musculus (strain C57BL/6). For each trial, 24-h starved adult female mosquitoes were provided with a blood meal from restrained live chicks or mice. Chicks were donated by the Bartolucci poultry farm (Córdoba, Argentina). Mice were commercially obtained at the Instituto de Investigación Médica Mercedes y Martín Ferreyra (CONICET - Universidad Nacional de Córdoba). Vertebrate hosts were offered 1 hour before the lights went out and kept for 3 hours during 2 consecutive gonotrophic cycles.

Revised version: For experimental trials, avian hosts (live chicks of the species Gallus gallus) and mammalian hosts (live mice of the species Mus musculus, strain C57BL/6) were used to evaluate the effect of blood meal source on fecundity, fertility, and hatchability. Chicks were generously donated by the Bartolucci poultry farm (Córdoba, Argentina), while mice were commercially obtained from the Instituto de Investigación Médica Mercedes y Martín Ferreyra (CONICET - Universidad Nacional de Córdoba). In each trial, 24-hour-starved adult female mosquitoes were provided with a blood meal from restrained chicks or mice. Vertebrate hosts were offered to mosquitoes one hour before the lights were turned off and were kept for 3 hours during 2 consecutive gonotrophic cycles.

Earlier version: The experimental use of animals (Gallus gallus and Mus musculus) was approved by the ethical committee at Facultad de Ciencias Médicas, Universidad Nacional de Córdoba (FCM-UNC) in compliance with the legislation regarding the use of animals for experimental and other scientific purposes (accession code: CE-2022-00518476-UNC-SCT#FCM).

Revised version: The experimental use of animals (G. gallus and M. musculus) was approved by the ethical committee at Facultad de Ciencias Médicas, Universidad Nacional de Córdoba (FCM-UNC) in compliance with the legislation regarding the use of animals for experimental and other scientific purposes (accession code: CE-2022-00518476-UNC-SCT#FCM).

Seasonality

Earlier version: To assess the effect of seasonality (photoperiod + temperature) on fecundity, fertility and hatchability, a “typical” summer and autumn day from Córdoba city were simulated in an incubator where mosquitoes were kept throughout the entire experiment. Summer conditions were: T°min = 22°C, T°max = 28°C, photoperiod = 14L:10D. Autumn conditions were: T°min = 16°C, T°max = 22°C, photoperiod = 10L: 14D. H umidity: 60-70% for both conditions. Data were obtained from Climate Data website (https://es.climate-data.org/)

Revised version: To assess the effect of seasonality (photoperiod + temperature) on fecundity, fertility, and hatchability, a “typical” summer and autumn day from Córdoba city was simulated in an incubator, where mosquitoes were housed throughout the entire experiment. Summer conditions were as follows: T°min = 22°C, T°max = 28°C, photoperiod = 14L:10D. Autumn conditions were characterized by: T°min = 16°C, T°max = 22°C, and photoperiod = 10L:14D. The humidity levels were maintained at 60-70% for both conditions. Data for simulated conditions were obtained from Climate Data website (https://es.climate-data.org/).

Feeding trial design

Added in the revised version: The interaction effect of blood source (bird or mammal) and seasonality (autumn and summer) was evaluated during two consecutive gonotrophic cycles (I and II). Eight experimental colonies were established using egg rafts collected from the maintenance colony: mammal-autumn-I (MA-I), mammal-summer-I (MS-I), bird-autumn-I (BA-I), bird-summer-I (BS-I), mammal-autumn-II (MA-II), mammal-summer-II (MS-II), bird-autumn-II (BA-II), and bird-summer-II (BS-II).

Earlier version: Four experimental colonies were established from egg rafts collected from the maintenance colony, based on the combination of blood source and seasonality: bird-summer, bird-autumn, mammal-summer and mammal-autumn. These experimental colonies were reared at controlled room conditions from eggs to adults to ensure all adults have the same size (Kauffman et al., 2017), then they were transferred to the incubator to begin the assay. Adults were kept in cardboard cages of 8.3 liters (21 cm x 24 cm), fed ad libitum with a 10% of sugar solution soaked in cotton pads on plastic glasses.

Revised version: The experimental colonies were maintained under controlled conditions to ensure that all adults were of the same size (Kauffman et al., 2017). Adult mosquitoes were placed in cardboard cages of 8.3 liters (21 cm x 24 cm) and were provided with ad libitum access to a 10% sugar solution, which was soaked in cotton pads placed on plastic cups.

Earlier version: Experiments were performed during two consecutive blood meals at 5 and 14 days post-emerging. After each blood meal, female mosquitoes were anesthetized with CO2 and classified as full engorged (VI Sella’s stage), partially fed (I-V Sella’s stage) and unfed (I Sella’s stage) (Silva Santos et al., 2019). Full engorged females were counted and separated into another cage to complete its gonotrophic cycle. Four days post-feeding a glass with distilled water was located inside the cage to allow oviposition. After each oviposition cycle egg rafts were collected, counted, and transferred to a 12-well plate with 3 mL of distilled water, where they were photographed to posteriorly count the number of eggs per raft. Rafts were maintained in each well until they hatched into L1 larvae, time where they were fixed with 1 mL of 96% ethanol. After hatching, the number of L1 larvae per raft was counted.

Revised version: Feeding trials were conducted at two time points: 5 days post-emergence (first cycle) and 14 days post-emergence (second cycle). Following each blood meal, female mosquitoes were anesthetized using CO2 and classified as fully engorged (VI Sella’s stage), partially fed (I-V Sella’s stage), or unfed (I Sella’s stage), following the classification by Silva Santos et al. (2019). Fully engorged females were counted and separated into a separate cage to complete their gonotrophic cycle. Four days after feeding, a cup containing distilled water was placed inside the cage to facilitate oviposition. After each oviposition cycle, egg rafts were collected, counted, and transferred to a 12-well plate containing 3 mL of distilled water. They were then photographed to subsequently determine the number of eggs per raft. Rafts were maintained in each well until they hatched into L1 larvae, at which point they were preserved with 1 mL of 96% ethanol, and the number of L1 larvae per raft was counted.

Statistical analysis

Added in the revised version: The reproductive outputs measured for each experimental colony were fecundity, fertility, and hatchability. All analyses were performed using R Studio statistic software, v.4.2.1 (R Core Team, 2022).

Added in the revised version: Fecundity was defined as the number of eggs per raft and fertility as the number of L1 larvae hatched per raft. To evaluate the effect of blood meal source and seasonality on these two variables, a Generalized Linear Model (GLM; package MASS) with negative binomial error distribution and logarithmic link function was adjusted (Venables & Ripley, 2002). Fecundity and fertility served as the response variables, while the explanatory variables included blood source, seasonality, and gonotrophic cycle, all in a three-way interaction.

Added in the revised version: Hatchability (= hatching rate) was defined as the ratio of the number of larvae to the number of eggs per raft. A GLM was applied using a quasipoisson error distribution and a logarithmic link function. The response variable was the number of larvae, while the explanatory variables included blood source, seasonality and gonotrophic cycle as interaction effects. Additionally, the logarithm of the number of eggs was incorporated as an offset in the model.

Added in the revised version: Multiple comparisons among treatments were conducted using the Tukey’s honestly significant difference test (HSD Tukey), incorporating the Kramer (1956) correction for unbalanced data, also known as the Tukey-Kramer method. This approach allows to set the family-wise error rate to 0.05. These contrasts were performed using the multcomp package (Hothorn et al., 2008).

Added in the revised version: Model validation was assessed using a simulation-based approach to generate randomized quantile residuals to test goodness of fit (Q-Q plot), homoscedasticity, and outliers and influential points (Cook’s distance and Leverage plots). These analyses were conducted with the DHARMa package (Hartig, 2022).

Added in the revised version: Model accuracy via simulations for the fecundity and fertility models

Added in the revised version: To assess the accuracy of the results, a simulation approach was conducted using a Monte Carlo algorithm to simulate data and model parameters according to the methodology proposed by Carsey & Harden (2013). For this purpose, parameters including coefficients, treatment means and dispersion parameter (referred to as theta, θ), were extracted from the fecundity and fertility models. Initially, a new dataset consisting of 1000 values (sample size of 125 per treatment) was generated. This was achieved using the rnbinom function from the MASS package, specifying the treatment means (computed from model coefficients) and the theta parameter. Subsequently, a negative binomial GLM was fitted to the simulated data, and the expected means for each of the eight treatments were obtained from the model. This simulation process was repeated 1000 times. The ratio between the number of iterations in which the simulated means fell within the observed 95% confidence interval divided by the total number of simulations was computed as an approximation of the robustness of results.

Removed in the revised version: The reproductive outputs measured for each experimental colony and gonotrophic cycle were fecundity, fertility and hatchability. All analyses were performed using R Studio statistic software, v.4.2.1 (R Core Team, 2022). Fecundity was defined as the number of eggs per raft and fertility as the number of L1 larvae hatched per raft. To evaluate the effect of blood meal source and seasonality on these two variables, a Generalized Linear Model (GLM; package MASS) with negative binomial error distribution and logarithmic link function was adjusted (Venables & Ripley, 2002), with fecundity and fertility as response variables, respectively. Blood source, seasonality and gonotrophic cycle in a three-way interaction were set as explanatory variables. Hatchability (= hatching rate) was defined as the ratio of the number of larvae to the number of eggs per raft. Because data did not follow any known exponential family probability distribution (needed for linear regression) a null model analysis via data randomization was performed. Prior to this analysis we adjusted a factorial ANOVA using hatchability as response variable and blood source, seasonality and gonotrophic cycle in a three-way interaction as explanatory variables. From this model we obtained and saved the observed F value (FO) for each variable (main effects) and its combination (interaction effects). Posteriorly, we randomly permuted the response variable 10000 times and obtained simulated F value to create an own distribution for the F parameter. The FO value was compared to the simulated distribution to test the null hypothesis that data were generated randomly against the alternative that there might be a pattern generated for at least one of the explanatory variables.

Results

Effect of blood source and seasonality on fecundity and fertility

Added in the revised version: Two interaction effects were found to be statistically significant in both the fecundity and fertility models (Tables 1-4). The first interaction effect was found between blood source and seasonality (fecundity: LRT X2 = 25.0, p = 7.72 x 10-7; fertility: LRT X2 = 13.82, p = 0.002), meaning that the effect of blood meal source on the number of eggs and larvae per raft varied with season (Tables 2, 3). In general, mosquitoes that fed on avian blood sources exhibited greater fecundity and fertility than those that fed on mammalian blood source, regardless of the season. When treatments were examined by pairwise comparisons, fecundity and fertility were 13% and 18% greater in autumn compared to summer for bird-fed mosquitoes, while they were 17% and 19% lower for mammal-fed mosquitoes (Fig. 1 A, C, Table 1). The second statistically significant interaction was between blood meal source and the gonotrophic cycle (fecundity: LRT X2 = 9.29, p = 0.0023; fertility: LRT X2 = 14.4, p = 0.00073). This indicated that the effect of blood type on the number of eggs and larvae per raft differed between the first and second feeding trial (Tables 2, 3). For mosquitoes that fed on mammalian blood, fecundity and fertility were 29% lower in the second cycle compared to the first. Although the difference was not statistically significant, fecundity and fertility showed a slight increase in the second cycle for bird-fed mosquitoes. During the second cycle, these outcomes were 41% and 40% higher, respectively, for bird-fed mosquitoes compared to mammal-fed mosquitoes (Fig. 1 B, D, Table 1).

Removed in the revised version: Blood meal source had a significant effect on fecundity but not in fertility. Seasonality had a significant effect on both fecundity and fertility. Gonotrophic cycle did not show an effect on neither of the two response variables (Tables 2, 3).

Removed in the revised version: It was observed an interaction effect between blood meal source and seasonality on fecundity and fertility (Table 2, 3). For both response variables, it was not found differences in the number of eggs and larvae per raft in summer for both blood meal sources, chicken and mouse. However, in autumn it was detected an increase in fecundity and fertility for bird-fed mosquitoes and a decrease for mouse-fed mosquitoes (Fig. 1 A, C, Table 1). In general, avian-fed mosquitoes presented greater fecundity and fertility than mammal-fed mosquitoes, regardless of the season. However, from summer to autumn, it was observed an increase of 13% and 18% on fecundity and fertility, respectively, for chicken-fed mosquitoes and a decrease of 17% and 19%, respectively, for mouse-fed mosquitoes.

Interaction plots showing the marginal effects of blood meal source-seasonality and blood meal source-gonotrophic cycle on (A, B) fecundity and (C, D) fertility. Predicted mean values and its corresponding confidence intervals are shown. Means sharing same letters are not statistically different.

Removed in the revised version: It was also detected an interaction effect between blood meal source and gonotrophic cycle (Table 2, 3). It was observed a decrease of 29% in fecundity and fertility from the first to second cycle for mouse-fed mosquitoes, being this treatment different from the remaining groups, in terms of fitness (Fig. 1 B, D, Table 1). During the second gonotrophic cycle fecundity and fertility for chicken-fed mosquitoes was 41% and 40% greater, respectively, than mosquitoes fed upon mice. It was not found statistically significant differences among the other three groups.

Removed in the revised version: The rest of single factors and interactions from the fecundity and fertility models were not statistically significant (Table 2, 3).

Effect of blood source and seasonality on hatchability

Added in the revised version: The overall hatchability reached 86%, with a range spanning from 76% for MA-II to a peak of 96% for MA-I, as outlined in Table 1 and depicted in Figure 2. The model found a statistically significant effect of blood meal source (LRT X2 = 73.06, p = 0.0012) and an interaction effect between seasonality and gonotrophic cycle (LRT X2 = 41.02, p = 0.0015) on hatching rates (Fig. 3, Table 4). Pairwise comparisons revealed a statistically significant 15% decrease in hatchability for mosquitoes from autumn from the first to the second gonotrophic cycle. Additionally, mosquitoes fed on mammals had greater hatchability than those fed on birds. In summer, hatchability during the first cycle was also higher than during the second cycle. These two trends were not found to be statistically significant.

Removed in the revised version: By visual inspection it was detected that mosquitoes fed upon mammals possessed the greater hatchability, except for mosquitoes feeding on mammals in autumn during the second gonotrophic cycle, which registered the least hatching rate (Fig. 2). Statistically, it was detected a significant effect of blood meal source on hatching rate (FO = 0.00074, p = 0.0004). Highest hatchability values were observed in mouse-fed mosquitoes, ranging from 82% (autumn, second gonotrophic cycle) to 96% (autumn, first gonotrophic cycle) (Table 1, Fig. 2).

Boxplots showing the effects of both blood meal source and seasonality on hatchability for the first and second gonotrophic cycle.

Added in the revised version: Interaction plot showing the marginal effects of seasonality-gonotrophic cycle on hatchability.

Added in the revised version: Model accuracy via simulations for the fecundity and fertility models

Added in the revised version: When comparing the simulated means for each of the eight treatments (combinations of blood meal source, seasonality, and gonotrophic cycle) with the means calculated from the observed data a high level of accuracy exceeding 90% was observed. Most of the simulated means fitted within the original 95% confidence interval, with an accuracy rate greater than 98%. The greatest mismatching among simulated and original data were found in the treatment MS-I, with an accuracy of 93% and 95% for the fecundity and fertility model, respectively. Accuracy remained high despite variations in sample sizes and data dispersion (Figs. 4, 5).

Added in the revised version: Simulated means of fecundity (eggs per raft) of each treatment along 1,000 iterations. Abbreviations. B: avian blood; M: mammalian blood; A: autumn; S: summer; I: first gonotrophic cycle; II: second gonotrophic cycle; Ac: model accuracy.

Added in the revised version: Simulated means of fertility (eggs per raft) of each treatment along 1,000 iterations. Abbreviations same as Figure 4.

Discussion

Added in the revised version: Our results revealed that there exists an interaction effect of blood meal source and seasonality on reproductive outputs of Cx. quinquefasciatus. Although we expected a higher fecundity and fertility in summer for bird-fed mosquitoes and higher outputs in autumn for mammal-fed mosquitoes, we found the opposite pattern.

Added in the revised version: Most studies have reported a greater fecundity and fertility in Cx. quinquefasciatus fed on different bird species (Akoh et al.,1993; Richards et al., 2012; Telang & Skinner,2019). This trend is also present in other species of Culex (Shroyer & Silvery, 1972) and Aedes mosquitoes (Harrison et al., 2021). The widely accepted hypothesis that has emerged to explain why avian blood produces a greater fecundity and fertility even in mosquitoes with a known mammal-preferred host choice focuses on the nutritional quality of host blood. Usually, it is addresses that because avian blood possesses nucleated and larger erythrocytes this offers greater quantities of proteins and energy than mammalian blood (Alto et al., 2014). However, this assumption lacks solid evidence due to studies regarding nutritional quality of avian blood and its impact in mosquito fitness are scarce. Actually, some studies have shown that mammalian blood are nutritionally richer than avian blood in terms of total protein content (Pearce 1977, Sant’Anna et al., 2010). Pre-diuresis, a physiological process where red blood cells and proteins are concentrated in the midgut by excreting large amounts of water, has been postulated as a mechanism for explaining the higher fecundity and fertility observed in bird-fed mosquitoes (Erram et al., 2022). Furthermore, avian blood has higher concentrations of isoleucine than mammalian blood, a crucial limiting amino acid necessary for egg production (Greenberg, 1951, Harrison et al., 2021).

Added in the revised version: The effect of temperature on fitness follows a unimodal relationship with a narrow optimal range. For mosquitoes of genus Culex and Aedes, this optimal band ranges from 20°C to 30°C, with an optimal at 28°C for species such as Cx. quinquefasciatus (Mordecai et al., 2019). Several studies have corroborated this trend, with higher fecundity and fertility towards the optimal of 28°C (Abouzied, 2017; Mogi, 1992), regardless of the blood source (bird or mammal). Regarding photoperiod, the effects are not as evident as temperature since data are somewhat contradictory. While Mogi (1992) found that mid and long day lengths induced greater fecundity, Costanzo et al. (2015) did not find differences of day length on fecundity. Because temperature and photoperiod are always coupled, in general their effects follow the same unimodal trend. Although our experiment was conducted within this optimal range, the highest fecundity and fertility in optimal conditions (summer) was present only in mammal-fed mosquitoes, while higher fitness in bird-fed mosquitoes was detected in autumn. This difference may indicate an interaction of seasonality with blood source.

Added in the revised version: While our findings may have certain limitations due to moderate sample sizes in some treatments and their reliance solely on laboratory experiments, they verify an interaction effect between blood meal source and seasonality, which had not been previously investigated or documented in any other study. Nevertheless, the precise reasons behind the observed pattern remain inconclusive. The highest levels of fecundity and fertility observed in bird-fed mosquitoes during autumn, when conditions are more challenging, may confer a reproductive advantage, leading to increased offspring production as an adaptive response to these adverse conditions. In contrast, during the more favorable conditions of summer, this pattern may not be as readily discernible. This proposed hypothesis needs further studies to be confirmed.

Added in the revised version: A second interaction was observed between blood meal source and gonotrophic cycle. We found that fecundity and fertility during the first cycle were similar between mosquitoes fed on chickens and those fed on mammals, and this was also similar with bird-fed mosquitoes from the second gonotrophic cycle, regardless of the season. The only difference we observed was between the fitness of mammal-fed mosquitoes from the second gonotrophic cycle and the earlier mentioned treatments. This pattern is not different from those found in other studies, as the first and second cycles are generally the most productive (Awahmukalah & Brooks, 1985; Christiansen-Jucht et al., 2015). Although we observed a slight increase of fecundity and fertility in bird-fed mosquitoes between cycles, it was statistically non-significant. To our knowledge, only two studies have reported this increasing pattern: an older one by Bennett (1970) and unpublished data from Obholz (personal communication), both conducted with Ae. aegypti. On the contrary, a decrease in fitness between cycles has been reported in other studies for Cx. quinquefasciatus (Richards et al., 2012; Telang & Skinner, 2019) fed upon birds or mammals. This is consistent with our findings regarding mouse-fed mosquitoes; however, because the fecundity and fertility of mouse-fed mosquitoes during the second cycle yielded the fewest egg rafts, this decreasing pattern might also be an effect due to the insufficient sample size.

Added in the revised version: In addition to assessing fecundity, we also investigated whether hatchability could account for the variations in fertility observed among different treatments. The only statistically significant trend we identified was a decrease in the percentage of hatched eggs from the first to the second gonotrophic cycle, but this was observed only in autumn mosquitoes. This decline in hatching rate may be attributed to a constant exposure to cooling conditions, due to the challenging environmental conditions during this season. This temperature-related effect aligns with findings from other studies (van der Linde et al., 1990). However, although our data suggest that changes in fertility may be linked to variations in hatching rate, it becomes apparent that fecundity has a more pronounced impact on larval production due to the effects of blood meal source and seasonality than hatching rate. Hence, the influence of hatchability on the number of larvae per raft is rather limited.

Added in the revised version: In summary, our study investigates the critical issue of host shifting patterns in Culex mosquitoes, which holds significant importance due to their role as bridge vectors between bird and human hosts. Our model suggests that changes in photoperiod and environmental temperature induce a different physiological state that requires the mosquito to feed on a blood meal source that results in higher fitness (increased fecundity and fertility). We predicted that mosquitoes fed on mammals in autumn will perform higher fecundity and fertility than those fed on birds in the same season. Unexpectedly, we observed the opposite pattern. Despite this, our study confirmed a significant interaction between blood meal source and seasonality, suggesting a complex interplay of factors affecting mosquito reproduction and its relationship with host use. Due to the observed opposite trend, we can consider three plausible explaining scenarios. Firstly, it is possible that additional factors, such as genetic mechanisms or variations in midgut microbiota composition, play a role in influencing host use patterns. Secondly, host switching might be explained by factors related to vector biology other than fitness. Lastly, it remains a possibility that Argentinian Cx. quinquefasciatus populations do not exhibit the same host-switching behaviour observed in US populations, and the observed interaction effects may be linked to different mosquito behaviours. Further investigations are warranted to corroborate our findings and gain a deeper understanding of the interaction between mosquito feeding pattern and seasonality. Ultimately, these investigations will provide valuable insights into mosquito-borne disease transmission and forecasting its emergence.

Removed in the revised version: Several host physiological factors have been studied as putative candidates to explain the changes in mosquito feeding behavior, including stress, health status or metabolic rate (Gervasi et al., 2016; Yan et al. 2017, 2018). However, it has not been investigated how seasonal use of hosts affects reproductive traits in mosquitoes. It is known that some factors like host blood source, temperature or photoperiod have an effect on some reproductive aspects (Takken & Verhulst, 2013). Some authors have studied and proved the significant effects of blood meal source (Telang & Skinner, 2019), temperature (Ciota et al., 2014) and photoperiod (Costanzo et al., 2015) on fecundity and fertility. These studies illustrate how single effects of blood source and environmental factors may influence reproductive outputs. However, to our knowledge, there are no available studies that have attempted to assess the interaction effect among these factors, and most articles tend to investigate the effects of blood meal source and temperature or photoperiod independently.

Removed in the revised version: Overall, we detected that chicken-fed mosquitoes produced more eggs and larvae per raft than mouse-fed mosquitoes, except in summer where differences were not detected. Akoh et al. (1993) reported a greater fecundity for mosquitoes fed on chicken than those fed on mice. Richards et al. (2012) and Telang & Skinner (2019) found a greater fecundity and fertility in chicken-fed mosquitoes than fed on bovine and pig blood, respectively. This effect is also present in other species of Culex (Shroyer & Silvery, 1972) and Aedes (Harrison et al., 2021). However, other authors have found the opposite pattern with human-fed mosquitoes (Ferdousi & Islam, 2005) and Guinea pig and human-fed mosquitoes (Xue et al., 2009) compared with mosquitoes fed upon chickens. In summary, with some few exceptions, mosquitoes fed on avian blood commonly produce a considerably greater batch size and fertility than mammal-fed mosquitoes. We also observed an increase of egg production and hatching per raft from summer to autumn in mosquitoes fed on chicks, whereas we found a decrease of these reproductive traits in mosquitoes fed upon mice. Regarding gonotrophic cycle, the second cycle was less productive in terms of fecundity and fertility for mosquitoes feeding on mice, however, we did not find differences in these reproductive outputs for mosquitoes fed on chicken. For hatchability, it was observed a greater hatching rate in mouse-fed mosquitoes than chicken-fed mosquitoes, however, the effect was not such marked as in fecundity and fertility.

Removed in the revised version: Several hypotheses (-I) nutritional hypothesis, -II) digestive physiology hypothesis and -III) digestive mechanical hypothesis) have been proposed to explain why avian blood produces a greater fecundity and fertility even in mosquitoes with a known mammal-preferred host choice. The first widely accepted general idea (I) suggests that since birds have larger and nucleated erythrocytes, they provide greater amounts of proteins and energy than the anucleated mammalian erythrocytes, inducing a higher fecundity and fertility in mosquitoes feeding from these animals (Alto et al., 2014). However, it has been shown that mammalian blood contains more total protein, hemoglobin and hematocrit levels than avian blood, suggesting it is nutritional richer. Even though, mosquitoes such as Cx. quinquefasciatus experiments a process known as pre-diuresis, excreting large amounts of water (derived from blood serum), allowing it to concentrate and accumulate more proteins. Hence, it could be possible that mosquitoes compensate the low protein quality of bird blood by consuming larger volumes or feeding for longer times (Erram et al., 2022) which might explain the greater fecundity and fertility while feeding on avian hosts. More than the total protein, it has been shown that it is the amino acid balance which determines the egg production, and amino acid content varies between different animals. For Aedes aegypti, eight amino acids are required to ensure egg production, namely, arginine, isoleucine, leucine, lysine, phenylalanine, threonine, tryptophan and valine. Consumption of these amino acids induces the expression of yolk protein precursor genes and the resulting synthesis of the protein vitellogenin, which is absorbed by oocytes and allows the development of embryos and posterior egg laying. It has been demonstrated that from all these amino acids, isoleucine is the limiting one, and Greenberg (1951) found that addition of isoleucine to washed sheep erythrocytes increases egg production and oviposition. However, this effect was not observed with the addition of a mixture of the other amino acids. Since mammalian blood is poor in isoleucine compared with avian blood, it was suggested and proven for several species that isoleucine-rich host blood induces a greater production of eggs by mosquitoes (Greenberg, 1951; Harrison et al., 2021). The digestive physiology hypothesis (II) states that differences in egg production between avian and mammalian hosts is due to variation in digestion rates. Downe et al. (1963) found that in Mansonia perturbans digestion of avian blood occurred more rapidly than mammalian blood and this pattern was also confirmed in Culex tarsalis (Downe and Archer, 1975). In Aedes aegypti, nutrients from a blood meal can enter the oocyte with a period of 36-48 hs, after that the pinocytotic activity of the gamete decreases and the intake of exterior material is blocked. Because of this, those species of mosquitoes that can achieve the digestion of blood in a short period will ensure the exploitation of nutrients generating a greater egg production. This could be the situation present in mosquitoes fed upon birds, which produce a greater fecundity (Downe & Archer, 1975). A third possible explanation is the digestive mechanical hypothesis (III). Mosquitoes need to retain erythrocytes in midgut and hemolyzed the blood meal to exploit proteins. Female mosquitoes are armed with a set of sclerotized teeth-like structures present in the anterior foregut (= cibarial armature) and posterior midgut (= pyloric armature), the first one acts as the primary site of hemolysis, whereas the second armature prevents blood to escape outside the midgut, enhancing the exploitation of resources from red blood cells. Since the structure of cibarial and pyloric armature varies between mosquitoes and because erythrocytes from host are of different shapes and sizes, these factors could contribute to mosquitoes to select those hosts whose blood is easier to hemolyze and its protein easier to assimilate (Lyimo & Ferguson, 2009; Vaughan et al., 1991). In summary, several factors regarding blood quality and digestive physiology make bird blood a higher nutritional resource for egg production, and despite these factors were usually studied independently, it must be the combination or interaction of all of them which enhance the production of a greater fitness in mosquitoes feeding on birds.

Removed in the revised version: Regarding environmental conditions, there is plentiful information available about the effects of temperature and photoperiod on physiological parameters (survival rate, development time, fecundity, fertility, behavior, etc.) (Hoffman, 1985). In nature, mosquitoes are exposed to a cyclic variation of both temperature and day length. It can be recognise a warm phase (= thermophase) which coincides with a photophase (light presence) and a cool phase (= cryophase) coinciding with a scotophase (light absence) (Joshi, 1996). Field and laboratory data have established that all these physiological traits follow a unimodal or hump-shaped relationship in ectothermic organisms, increasing from zero at a thermal minimum to a highest optimum and then decreasing to zero again at a thermal maximum. The best performance of animals generally is obtained at a narrow “optimal range”, which varies among and within species and traits. In Culex quinquefasciatus the fecundity curve ranges from a thermal minimum of 5°C to a maximum of 37.7°C with an optimal temperature of 21.4°C (Mordecai et al., 2019). However, the greatest egg production is restricted to a narrow interval around the thermal optimum. This narrow band has been determined for some species, for Culex pipiens complex and Aedes aegypti it seems that the optimal temperature ranges from 20°C to 30°C (Spanoudis et al., 2019; Bader et al. 2012).

Removed in the revised version: The effect of temperature on physiological traits has been studied in two ways, keeping it constant across all the experiments or fluctuating it in different ways, generally between a minimum and maximum. For different populations of Culex quinquefasciatus and for a constant temperature, Mogi (1992) has found that fecundity increases from 15°C to 20°C and decreases to 28°C. Same pattern was registered for Aedes albopictus, with a peak at 26°C and decreasing to 30°C. Fluctuating temperature has the same pattern, although because temperature is allowed to fluctuate within a range mosquitos seem to cope better the extreme conditions, for example, Aedes albopictus produce same quantity of eggs at 18°C-26°C and 22°C-30°C (Joshi, 1996). Regarding photoperiod, fewer studies are available, some of them are those of Mogi (1992) and Costanzo et al. (2015), however, the effects are not as evident as temperature, and while Mogi (1992) found that mid and long day lengths induced greater fecundity, Costanzo et al. (2015) did not find differences of day length on fecundity, then, data are more contradictory than the effects of temperature.

Removed in the revised version: Two important mosquito traits that might affect fecundity and the subsequent fertility are calendar age (= age post-emerging) and physiological age (= gonotrophic cycle). In general, fecundity and fertility decline with calendar age (Akoh et al., 1993; McCann et al., 2009) and gonotrophic cycle (Richards et al., 2012; Telang & Skinner, 2019). This general pattern was shown by us in mouse-fed mosquitoes. Although we also recorded a numerical increase in fecundity and fertility between gonotrophic cycles for mosquitoes fed on chicken, this difference was not significant. Same pattern was found by Bennet (1970) in Aedes aegypti fed upon avian and mammalian hosts. In addition, Obholz (unpublished data) also found an increase in fecundity and fertility for Ae. Aegypti fed on mice. A difference in egg production may be attributable to variation in feeding rate, but since we only evaluated full-engorged female mosquitoes this difference cannot be explained based on this rate. Hence, there must be an underlying physiological or environmental process that may be affecting fecundity and fertility and further studies are needed to corroborate it.

Removed in the revised version: All studies assessing the effect of temperature or photoperiod on reproductive traits have been independent from the source of blood meal, however, our results suggest that these two environmental variables have a significant effect on fecundity and fertility, affecting differentially regarding blood meal source. We found that in summer (mean temperature = 25°C and long-day photoperiod) there is not differences on fecundity and fertility for mosquitoes fed upon mouse and chicken, however, in autumn (mean temperature = 19°C and short-day photoperiod) the fecundity and fertility of chicken-fed mosquitoes increased whereas in mouse-fed mosquitoes decreased. These results are not directly comparable with literature, since there are no studies that have evaluated the interaction effect of blood meal source and seasonality (temperature + photoperiod). Several authors suggest that fecundity and fertility increase at higher temperatures and photoperiod for both bird and mammal-fed mosquitoes (Abdallah et al., 2020; Al-Rashidi et al., 2022, Shehata, 2018), however, a difference based on blood meal source is elusive. This little evidence makes it evident that more studies are needed to assess the combined effect of temperature, photoperiod and blood source on reproductive traits since a common pattern has not been readily identified at the moment.

Removed in the revised version: Fertility in our experiments varied among treatments. Considering the number of eggs, differences in fertility may be attributable to two factors: (I) variation in fecundity or (II) variation in hatching rate (= hatchability). The first has already been tested, then, to assess if the variation in number of larvae was a result of differences in hatchability, we evaluated if there were changes in the percentage of eggs hatched. We find significant differences in hatchability regarding the type of blood meal source, obtaining a mean hatching rate for mammal-fed mosquitoes of 92% and 87% for bird-feeding mosquitoes. In literature, hatchability (also cited as viability or even fertility) has been shown to vary according to different temperatures (van der Linde et al., 1990) but not with blood meal sources (dos Santos Dias et al., 2018; Richard et al., 2012). Our preliminary data suggest that changes in fertility are a consequence of variation in egg production due to the effects of blood meal source and seasonality and not in hatchability

Removed in the revised version: Understanding factors leading host selection and host switch are critical to predict arboviral outbreaks and arbovirus spill over. It is well known that many species of Culex mosquitoes experience changes in their host usage, from birds in spring and late summer to mammals in late summer and autumn, which usually trigger the increase of viral infection reported on humans, turning from a period of viral amplification to viral transmission. Nevertheless, studies considering the causes of host switching are very scarce and they mainly focus on host factors, such as host availability or defensive behavior (Kilpatrick et al., 2006, Burkett-Cadena et al., 2011). A variety of mosquito physiological traits have been studied as an attempt to explain host selection but neither of them have been applied in host switch mechanisms (Lyimo & Ferguson, 2009; Yan et al., 2017, 2018). In this work we investigated for the first time if some reproductive traits could explain the pattern in host shift from birds to mammals. In our study we included the interaction effects of blood meal source and seasonality on fecundity and fertility, two variables analyzed separately but never combined, to the best of our knowledge. Our results suggest that there is a change in fecundity and fertility through seasons due to the effects of blood meal source and seasonality.

Removed in the revised version: In summary, our model suggests that biological and climatic factors induce physiological changes in mosquitoes which triggers selection of new hosts, generating an increase of fitness (fertility and fecundity). Our results did not support the hypothesis that greater fitness on birds is present in summer and on mammals in autumn, since we register the opposite pattern. Notwithstanding, since Culex mosquitoes have a wide range of host niche it remains inconclusive whether other factors may be intervening in mosquitoes seasonal host shift.